Ligand-dependent Enhancer Activation Indirectly Modulates Non-target Promoters in a Chromatin Domain

Author response:

The following is the authors’ response to the previous reviews

We are pleased that Reviewer 3 appreciated our findings and found the temporal lag between the expression of TFF1 and TFF3 during signaling particularly interesting. The reviewer also advised us not to overemphasize that this lag arises from phase separation of ERα at the TFF1 locus, as the use of 1,6-hexanediol alone is not sufficient to conclusively establish whether ERα condensates undergo liquid–liquid phase separation.

We agree with this assessment and have revised the manuscript accordingly. Specifically, we have modified the title to remove reference to phase separation and have updated the text throughout the manuscript to avoid claiming that the observed condensates are a result of phase separation.

The revised title is:

“Ligand-dependent Enhancer Activation Indirectly Modulates Non-target Promoters in a Chromatin Domain.”

With these changes, we are proceeding with the version of record using revised version of the manuscript.

Thank you for your continued support.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Transcription activation of genes by estrogen is driven by enhancers, which are often located within the same Topologically Associating Domain (TAD) as non-targeted promoters. We investigated how acute enhancer-driven activation affects neighbouring non-target genes within the same TAD. Using single-molecule RNA FISH (smFISH), we tracked the transcription of TFF1 (enhancer-targeted) and TFF3 (non-targeted) during estrogen stimulation. We observed mutually exclusive expression patterns: TFF1 expression peaked at 1 hour, while TFF3 reached its peak at 3 hours, after TFF1’s activation had diminished. Chromatin looping data indicated that the enhancer loops with TFF1 but not TFF3, suggesting that TFF3 upregulation is not due to direct enhancer-promoter interactions. CRISPR deletion of the TFF1 enhancer and 1,6-hexanediol (HD) exposure revealed that the TFF1 enhancer:promoter undergo Liquid-Liquid Phase Separation (LLPS), which sequesters the transcriptional machinery and inhibits TFF3 expression. As estrogen signalling wanes or LLPS is disrupted, TFF1 expression declines while TFF3 expression increases. Our findings reveal that enhancer-driven activation can indirectly influence neighbouring genes, highlighting a dynamic shift in gene expression as signalling progresses.

Revised version: Transcription activation of genes by estrogen is driven by enhancers, which are often located within the same Topologically Associating Domain (TAD) as non-targeted promoters. We investigated how acute enhancer-driven activation affects neighbouring non-target genes within the same TAD. Using single-molecule RNA FISH (smFISH), we tracked the transcription of TFF1 (enhancer-target gene) and TFF3 (non-target gene) during estrogen stimulation. We observed mutually exclusive expression patterns: TFF1 expression peaked at 1 hour, while TFF3 reached its peak at 3 hours, after TFF1 activation had diminished. Chromatin looping data indicated that the enhancer loops with TFF1 but not TFF3, suggesting that TFF3 upregulation is not due to direct enhancer-promoter interactions. CRISPR deletion of the enhancer, affected TFF1 transcription more acutely than TFF3. 1,6-hexanediol (HD) exposure suggested that the TFF1 enhancer:promoter undergo a potential ERα-mediated condensate formation, which sequesters the transcriptional machinery and inhibits TFF3 expression. As estrogen signalling fades at 3h, TFF1 expression declines while TFF3 expression increases. Our findings reveal that enhancer-driven activation can indirectly repress neighbouring genes within the same TAD, highlighting a dynamic shift in gene expression as signalling progresses.

Introduction

Acute transcriptional activation of ligand-induced genes drives downstream signalling response. Similar to development-specific genes, signalling induced genes are also driven by enhancers (Hah et al., 2013; Li et al., 2013; Liu et al., 2014; Uyehara & Apostolou, 2023). Upon binding with ligand-induced transcription factors (TFs), these enhancers loop with their target promoters mostly, in the same TAD (Buecker et al., 2014; Bulger & Groudine, 1999; Chepelev et al., 2012; Furlong & Levine, 2018; Oh et al., 2021; Panigrahi & O’Malley, 2021; Ptashne, 1986; Sanyal et al., 2012; Yan et al., 2018). Enhancer: promoter pairing is thought to be specific and forms the basis of noise-free gene activation of a subset of genes crucial for signalling response (Bojcsuk et al., 2017; H. Chen et al., 2018; Friedman et al., 2024; Galouzis & Furlong, 2022; Zabidi et al., 2015). Due to such specificity, gene transcription occurs in waves of early and late responsive genes (Fowler et al., 2011; Yamamoto & Alberts, 1976). Often, the protein factors translated from early genes regulate the expression of late responsive genes (Dixon et al., 1996; Freter et al., 1996; Herschman, 1991; Williams et al., 1999; Winkles, 1997; Winston & Pledger, 1993). However, it is not known if the genes that are activated early in the signalling time course are spatially related to late activating genes. Further, if the spatial proximity of any gene to an early gene is enough to cause the temporal differential transcription due to sequestration of the transcriptional machinery from late gene to early gene is poorly understood. Transcription factors (TFs), after binding to cognate DNA motifs, recruit RNA polymerase machinery and co-activators for gene activation (Levine & Tjian, 2003; Liu et al., 2014). These machinery are limited in supply, and may become sequestered from other genomic regions (Koşar & Erbaş, 2022). Further, acute activation of genes is linked with phase-separation of TFs, polymerases, mediators and other co-factors/activators, most of which harbour low complexity regions to promote weak protein-protein interactions driving the formation of TF-condensates (Boehning et al., 2018; Boija et al., 2018; Cai et al., 2019; L. Chen et al., 2023; Cho et al., 2018; Chong et al., 2018, 2022; Mann & Notani, 2023; Sabari et al., 2018; Shrinivas et al., 2019; Stortz et al., 2020, 2024). Though the precise stoichiometry of protein molecules in these condensates is unknown, such structures involve multiple molecules of each transcriptional protein. The formation of such phase-separated compartments can potentially act as a sink for transcriptional machinery depriving neighbouring promoters and enhancers that are not part of the compartment. Such sequestration would cause indirect suppression of these spatially proximal genes in the same TAD but not the part of the same condensate.

In order to investigate the effect of an enhancer: promoter pair on a neighbouring gene within the same TAD, we looked at the paradigmatic model of estrogen signalling in mammary epithelial cells, namely MCF7 (Levenson & Jordan, 1997; Masiakowski et al., 1982). E2-signalling can be induced within minutes by treating the cells with estradiol (E2), which causes activation and repression of several genes across the genome. The peak of E2-mediated signalling occurs 40 minutes post-induction and starts to decay by 160 minutes (Hah et al., 2011). This occurs via the binding of estrogen receptor-alpha (ERα) on the enhancers (Li et al., 2013). We selected E2-induced E-P pair TFF1 and its neighbouring gene TFF3 on Chr21 as discussed below (Chinery et al., 1996). Several studies have used single-cell methods to interrogate E2 responsive genes such as GREB1, MYC and TFF1 as a paradigm to understand their transcriptional response upon estrogen stimulation (Fritzsch et al., 2018; Patange et al., 2022; Rodriguez et al., 2019; Stossi et al., 2020). These studies have implicated the cellular states as determinants of transcriptional response and that the long repressive states of TFF1 give rise to expression variability in isogenic lines (Rodriguez et al., 2019).

Earlier version: TFF1 is found in a topologically associated domain (TAD) along with a few other genes, including TFF2, TFF3, TMPRSS3, and UBASH3A (Oh et al., 2021; Quintin et al., 2014; Rao et al., 2014). TFF1 and TFF3 is located at 10kb and 60kb from enhancer, respectively (Fig. 1A), and this locus has been useful in answering the questions about promoter-enhancer interactions and multi-gene regulation (Oh et al., 2021; Quintin et al., 2014;).

Revised version: TFF1 is found in a topologically associated domain (TAD) along with a few other genes, including TFF2, TFF3, TMPRSS3, and UBASH3A (Oh et al., 2021; Quintin et al., 2014; Rao et al., 2014). TFF1 and TFF3 is located at 10kb and 60kb from enhancer, respectively (Figure 1A), and this locus has been useful in answering the questions about promoter-enhancer interactions and multi-gene regulation (Oh et al., 2021; Quintin et al., 2014;).

Earlier version: ERa binding and ligand-induced gene expression of TFF1 and TFF3 change over the course of estrogen signaling

Revised version: ERα binding and ligand-induced gene expression of TFF1 and TFF3 change over the course of estrogen signaling

Removed in the revised version: A. Schematic depicting TFF1 locus, UCSC genome browser snapshots showing the binding of ERα, H3K27ac status, H3K4me3 signal, and Gro-seq signal for robustly E2-induced TFF1 locus. First, second and third ERα ChIP-seq and Gro-seq tracks are from vehicle-treated, E2-1h and E2-3h in WT cells, respectively.

Earlier version: B. qRT-PCR showing the changes in expression of TFF1 and TFF3 genes during the E2 signaling time course. Error bars denote SEM from four biological replicates. Each dot represents a replicate. p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Revised version: A. Schematic depicting TFF1 locus, UCSC genome browser snapshots showing the binding of ERα, H3K27ac status, H3K4me3 signal, and Gro-seq signal for robustly E2-induced TFF1 locus. First, second and third ERα ChIP-seq and Gro-seq tracks are from vehicle-treated, E2-1h and E2-3h in WT cells, respectively. B. qRT-PCR showing the changes in expression of TFF1 and TFF3 genes during the E2 signaling time course. Error bars denote SEM from four biological replicates. Each dot represents a replicate. p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Earlier version: We chose to answer these questions using an integrated approach consisting of single molecule RNA FISH (smFISH), genome-wide conformation capture and sequencing (4C-seq), and perturbation of cis-acting regulatory elements by CRISPR. We identified that TFF1 and TFF3 genes located within the same TAD show distinct and opposing expression profiles over the course of E2-mediated signalling. In agreement with this, the enhancer also showed increased lopping interaction with the cognate gene promoter during peak expression at 1h of signalling. We also identified the role of TF-mediated LLPS on the regulation of these two genes during signalling. Hence, we propose that ligand-dependent condensation of ERα can support the expression of the cognate gene in concert with the enhancer while negatively affecting the expression of a neighbouring gene.

Revised version: We chose to answer these questions using an integrated approach consisting of single molecule RNA FISH (smFISH), genome-wide conformation capture and sequencing (4C-seq), and perturbation of cis-acting regulatory elements by CRISPR. We identified that TFF1 and TFF3 genes located within the same TAD show distinct and opposing expression profiles over the course of E2-mediated signalling. In agreement with this, the enhancer also showed increased lopping interaction with the cognate gene promoter during peak expression at 1h of signalling. We identified a potential role of TF-mediated condensates on the up-regulation of TFF1 and down-regulation of TFF3 at the peak of signalling. Hence, we propose that ligand-dependent ERα-protein assemblies can support the expression of the cognate gene in concert with the enhancer that indirectly represses expression of a neighbouring gene in the same TAD.

Results

ERα binding and ligand-induced gene expression of TFF1 and TFF3 change over the course of estrogen signaling

Earlier version: In order to understand how acute activation of one gene in a TAD affects the transcription of a proximate gene, we chose to study TFF1 and its neighbouring gene TFF3 at a 43kb distance within same TAD. TFF1 expression is linked with an enhancer located 10 kb downstream (Li et al., 2013; Saravanan et al., 2020; Oh et al., 2021) and the distance between TFF1 enhancer and TFF3 gene is 53kb (Fig. 1A).

Revised version: In order to understand how acute activation of one gene in a TAD affects the transcription of a proximate gene, we chose to study TFF1 and its neighbouring gene TFF3 at a 43kb distance within the same TAD. TFF1 expression is linked with an enhancer located 10 kb downstream (Li et al., 2013; Saravanan et al., 2020; Oh et al., 2021) and the distance between TFF1 enhancer and TFF3 gene is 53kb (Figure 1A).

Earlier version: 17-β-oestradiol exposure leads to ERα binding on regulatory regions, leading to gene activation. Gene transcription of E2-regulated genes was shown to peak at 1h and significantly reduced at 3h due to the rapid degradation of ERα (Hah et al., 2011). The binding of ERα in the genome also follows this temporal kinetics, where it peaks at 1h and reduces at 3h post ligand stimulation (Hah et al., 2011; Li et al., 2013; Liu et al., 2014). At the TFF1 locus, the TFF1 enhancer and to some extent, its promoter, were bound by ERα even in the absence of E2 induction, while it increased in strength at these regions and also at various other regions in the locus at 1h of ligand stimulation (Saravanan et al., 2020). Notably, the binding at these regions substantially decreased at 3h. On the other hand, ERα did not bind on the TFF3 promoter throughout the course of signalling (Fig. 1A and Fig. S1B). A single peak of ERα was observed approx. 3 kb upstream to TFF3 gene. We have previously shown that estrogen-induced clustered binding of ERα on TFF1 region is associated with its acute activation post estrogen stimulation (Saravanan et al., 2020). We then tested their expression level by nascent RNA-seq (GRO-seq). We observed a dramatic transcriptional activation of TFF1 at 1h, which reduced substantially at 3h (Fig. 1A). The expression level of TFF3 was comparatively lower and exhibited mild induction at 1h which increased furthermore at 3h upon signalling (Fig. 1A, Fig. S1C). However, qRT-PCR on unspliced TFF3 exhibited upregulation at 1h and down-regulation at 3h (Fig. 1B). We reasoned that these differences in TFF3 could be because of its low baseline expression (Clark et al., 2015; Conesa et al., 2016; Sha et al., 2015; Svensson et al., 2017; Tarazona et al., 2011), and lower yet expression of unspliced transcripts. Therefore, in order to test the relative induction of these genes and their co-regulation, we decided to perform single-cell measurements of RNA (Fig. S1A) that do not rely on PCR amplification and can provide absolute numbers of TFF1 and TFF3 transcripts in a given cell. Towards this, we chose to employ smFISH (Femino et al., 1998; Haimovich & Gerst, 2018; Kwon, 2013; Raj et al., 2008).

Revised version: 17-β-oestradiol exposure leads to ERα binding on regulatory regions, leading to gene activation. Gene transcription of E2-regulated genes was shown to peak at 1h and significantly reduced at 3h due to the rapid degradation of ERα (Hah et al., 2011). The binding of ERα in the genome also follows this temporal kinetics, where it peaks at 1h and reduces at 3h post ligand stimulation (Hah et al., 2011; Li et al., 2013; Liu et al., 2014). At the TFF1 locus, the TFF1 enhancer and to some extent, its promoter, were bound by ERα even in the absence of E2 induction, while it increased in strength at these regions and also at various other regions in the locus at 1h of ligand stimulation (Saravanan et al., 2020). Notably, the binding at these regions substantially decreased at 3h. On the other hand, ERα did not bind on the TFF3 promoter throughout the course of signalling (Figure 1- figure supplement 1B). A single peak of ERα was observed approx. 3 kb upstream to TFF3 gene. We have previously shown that estrogen-induced clustered binding of ERα on TFF1 region is associated with its acute activation post estrogen stimulation (Saravanan et al., 2020). We then tested their expression level by nascent RNA-seq (GRO-seq). We observed a dramatic transcriptional activation of TFF1 at 1h, which reduced substantially at 3h (Figure 1A). The expression level of TFF3 was comparatively lower and exhibited mild induction at 1h which increased furthermore at 3h upon signalling (Figure 1- figure supplement 1B). However, qRT-PCR on unspliced TFF3 exhibited upregulation at 1h and down-regulation at 3h (Figure 1B). We reasoned that these differences in TFF3 could be because of its low baseline expression (Clark et al., 2015; Conesa et al., 2016; Sha et al., 2015; Svensson et al., 2017; Tarazona et al., 2011), and lower yet expression of unspliced transcripts. Therefore, in order to test the relative induction of these genes and their co-regulation, we decided to perform single-cell measurements of RNA (Figure 1- figure supplement 1A) that do not rely on PCR amplification and can provide absolute numbers of TFF1 and TFF3 transcripts in a given cell. Towards this, we chose to employ smFISH (Femino et al., 1998; Haimovich & Gerst, 2018; Kwon, 2013; Raj et al., 2008).

TFF1 and TFF3 exhibit opposite trends during the E2 signalling time-course

MCF7 cells are hypertriploid to hypotetraploid, and each of the three alleles within a nucleus can behave differently, and smFISH is pre-eminently suited to capturing cell- and allele-specific heterogeneity of expression even for low expressing genes, compared to bulk studies that average over cell populations. Briefly, smFISH allows to visualize single RNA molecules in fixed cells using multiple fluorescently labelled oligonucleotide probes targeted to the RNA of interest and therefore, it can be used to image the transcription and the localization of multiple gene transcripts at same and different time points after signalling (Femino et al., 1998; Raj et al., 2008).

We designed the smFISH probes targeting the intronic region of TFF1 and TFF3 to measure their nascent transcripts (referred to as InTFF1 and InTFF3 hereon) whereas the exonic probes were used to primarily measure mature mRNA transcripts (referred to as ExTFF1 and ExTFF3 hereon). Intronic probes are particularly suited for investigating nascent transcriptional status at the time of fixation, while probes against mature mRNA reflects on more steady-state levels of functional mRNA due to finite mRNA lifetimes (Skinner et al., 2016). Additionally, the intronic probes can help determine the localisation and number of alleles that are transcribing at a given time thus allowing for additional interpretations regarding the expression of multiple genes (Skinner et al., 2016).

Earlier version: Representative images from the smFISH experiment probing InTFF1 and InTFF3 are depicted (Fig. 2A). The smaller foci represent individual intronic/nascent transcripts while the larger foci represent the site of transcription (Raj et al., 2006; Zenklusen et al., 2008). Since MCF7 cells are hypertriploid in nature, we expected to see 1-3 large foci per cell representing sites of transcription. Additionally, we observed many individual transcripts labelled by intronic probes (Fig. 2A, D) within the same nuclei (probes against mature RNA are usually more cytoplasmic). This is suggestive of the fact that the transcripts undergo non-co-transcriptional splicing as the transcripts labelled by intronic probes are localised away from the site of transcription. This is not surprising as several reports have shown that nascent transcripts can undergo splicing well after transcription (Coulon et al., 2014; Drexler et al., 2020; Khodor et al., 2012).

Revised version: Representative images from the smFISH experiment probing InTFF1 and InTFF3 are depicted (Figure 2A). The smaller foci represent individual intronic/nascent transcripts while the larger foci represent the site of transcription (Raj et al., 2006; Zenklusen et al., 2008). Since MCF7 cells are hypertriploid in nature, we expected to see 1-3 large foci per cell representing sites of transcription. Additionally, we observed many individual transcripts labelled by intronic probes (Figure 2A, D) within the same nuclei (probes against mature RNA are usually more cytoplasmic). This is suggestive of the fact that the transcripts undergo non-co-transcriptional splicing as the transcripts labelled by intronic probes are localised away from the site of transcription. This is not surprising as several reports have shown that nascent transcripts can undergo splicing well after transcription (Coulon et al., 2014; Drexler et al., 2020; Khodor et al., 2012).

TFF1 and TFF3 expressions show opposite trends during the E2 signaling time-course

Removed in the revised version: A. 60X Representative images from single molecule RNA FISH experiment showing transcripts for TFF1 and TFF3. The probe was designed against the unspliced RNA containing the intronic region. The scale bar is 5 microns.

Removed in the revised version: B. The mean RNA numbers are depicted. These are counted using an in-house MATLAB code which uses the DAPI-stained nuclei as the mask to count the RNA present in the nucleus. The graph shows the mean of means from three different repeats of the experiment, and error bars denote SEM (n= 665, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: C. Scatter plots showing the distribution of InTFF1 and InTFF3 on a cell-by-cell basis (n= 665, N=3). The absolute RNA numbers are combined from three different repeats. Density plots have been used to clearly visualize overlapping data points.

Removed in the revised version: D. 60X Representative images from single molecule RNA FISH experiment showing transcripts for InTFF1 and ExTFF1. Scale bar is 5 microns.

Earlier version: E. The mean RNA numbers for InTFF1 and ExTFF1 are depicted. Separate probes were used to target unspliced (InTFF1) and mature (ExTFF1) RNA. These are counted using an in-house MATLAB code which uses the DAPI-stained nuclei as the mask to count the intronic RNA present in the nucleus and a free-drawn region to designate the cell to count the exonic RNA present in the nucleus as well as the cytoplasm. The graph shows the mean of means from three different repeats of the experiment, and error bars denote SEM (n>360, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Revised version: A. 60X Representative images from single molecule RNA FISH experiment showing transcripts for TFF1 and TFF3. The probe was designed against the unspliced RNA containing the intronic region. The scale bar is 5 microns. B. The mean RNA numbers are depicted. These are counted using an in-house MATLAB code which uses the DAPI-stained nuclei as the mask to count the RNA present in the nucleus. The graph shows the mean of means from three different repeats of the experiment, and error bars denote SEM (n= 665, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. C. Scatter plots showing the distribution of InTFF1 and InTFF3 on a cell-by-cell basis (n= 665, N=3). The absolute RNA numbers are combined from three different repeats. Density plots have been used to clearly visualize overlapping data points. D. 60X Representative images from single molecule RNA FISH experiment showing transcripts for InTFF1 and ExTFF1. Scale bar is 5 micrometers. E. The mean RNA numbers for InTFF1 and ExTFF1 are depicted. Separate probes were used to target unspliced (InTFF1) and mature (ExTFF1) RNA. These are counted using an in-house MATLAB code which uses the DAPI-stained nuclei as the mask to count the intronic RNA present in the nucleus and a free-drawn region to designate the cell to count the exonic RNA present in the nucleus as well as the cytoplasm. The graph shows the mean of means from three different repeats of the experiment, and error bars denote SEM (n>360, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. F. The mean RNA numbers for InTFF3 and ExTFF3 are depicted. Separate probes were used to target unspliced (InTFF3) and mature (ExTFF3) RNA. The graph shows the mean of means from three different repeats of the experiment, and error bars denote SEM (n>210, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. G. Violin plots showing the ratio of intronic to exonic TFF1 counts are depicted. The graph shows the distribution of ratios combined from three different repeats (n>360, N=3). p-values were calculated by the Mann-Whitney test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. H. Violin plots showing the ratio of intronic to exonic TFF3 counts are depicted. The graph shows the distribution of ratios combined from three different repeats (n>210, N=3). p- values were calculated by the Mann-Whitney test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: F. The mean RNA numbers for InTFF3 and ExTFF3 are depicted. Separate probes were used to target unspliced (InTFF3) and mature (ExTFF3) RNA. The graph shows the mean of means from three different repeats of the experiment, and error bars denote SEM (n>210, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: G. Violin plots showing the ratio of intronic to exonic TFF1 counts are depicted. The graph shows the distribution of ratios combined from three different repeats (n>360, N=3). p-values were calculated by the Mann-Whitney test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: H. Violin plots showing the ratio of intronic to exonic TFF3 counts are depicted. The graph shows the distribution of ratios combined from three different repeats (n>210, N=3). p-values were calculated by the Mann-Whitney test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Earlier version: We quantified the number of such transcripts per nuclei as a proxy for ongoing transcription. At 1h post E2 induction, TFF1 mean transcript counts increased significantly compared to uninduced, whereas, the increment was less pronounced for TFF3 transcripts. In contrast, the mean transcript counts for TFF3 increased significantly at 3h post induction while TFF1 transcription showed a decrease at 3h compared to 1h (Fig. 2B). We observed an increase in the number of transcribing sites per cell for both TFF1 and TFF3, as well as in the number of cells showing transcripts compared to the untreated condition (Fig. 2B, Fig. S2A-B). Further, TFF1 transcripts spots were more spread out in the nucleus that harboured multiple transcribing sites (Fig. S2C-D). To check the transcriptional status of TFF1 and TFF3 in the same cell, we plotted the transcript counts from individual cells (Fig. 2C). The data suggested that the transcript counts for TFF1 increased at 1h and cells that showed high counts for TFF1 were likely to have low counts for TFF3 and vice-versa. But the RNA counts for TFF1 decreased at 3h and increased for TFF3. Overall, it was evident that the transcriptional profile for these two genes located in the same TAD were negatively correlated as they peaked at different time points. To further confirm this, smFISH using probes targeting both the intronic and exonic transcripts in the same experiment was conducted. Intronic probes represent active transcription, while exonic probes show accumulated mature RNA from past transcriptional events even in the absence of active transcription (Skinner et al., 2016). Representative images from the smFISH experiment probing InTFF1/ExTFF1 and InTFF3/ExTFF3 are depicted (Fig. 2D, Fig. S2E). InTFF1 spots overlapped with ExTFF1 in the nucleus (Fig. S2F) and InTFF1 spots were restricted to the nucleus, whereas ExTFF1 were both in the nucleus and the cytoplasm as expected, indicating the specificity of smFISH probes (Fig. S2E). We observed that the exonic transcript counts for both TFF1 and TFF3 increased at 1h compared to uninduced (Fig. 2E, F). Strikingly, at 3h, the exonic transcripts for TFF1 continued to increase even while the intronic transcript counts reduced, though the fold increase (1.12 ±0.02) was less compared to that between uninduced and 1h (1.63 ± 0.28), indicating a plateauing of steady state levels. Exonic transcripts for TFF3 at 3h remained comparable to 1h while the intronic counts increased suggesting, TFF1 transcription reduces at 3h, whereas TFF3 expression increases. We reasoned that the ratio of intronic transcript number (InTFF) to exonic transcripts number (ExTFF) should represent the status of transcription as an increase in the number of intronic transcript due to expression would result in a higher ratio while also taking into consideration the number of mature transcripts. As expected, the ratio of intronic transcripts to exonic transcripts also showed that transcription is active at 1h for TFF1 as the ratio is higher compared to uninduced and 3h (Fig. 2G and Fig. S3A). Contrastingly, the ratio was highest at 3h for TFF3 suggesting that active transcription takes place much after the peak of E2-mediated signalling and maximal TFF1 expression (Fig. 2H and Fig. S3B). Indeed, a very recent study has shown that such post-transcriptional splicing can occur for genes and intron dispersal can be expected more commonly for some highly expressed genes (Coté et al., 2023). Therefore, the observation of intronic signal away from the site of transcription as we see for TFF1 and TFF3 is not unexpected (Fig. 2A, D, S2).

Revised version: We quantified the number of such transcripts per nuclei as a proxy for ongoing transcription. At 1h post E2 induction, TFF1 mean transcript counts increased significantly compared to uninduced, whereas, the increment was less pronounced for TFF3 transcripts. In contrast, the mean transcript counts for TFF3 increased significantly at 3h post induction while TFF1 transcription showed a decrease at 3h compared to 1h (Figure 2B). We observed an increase in the number of transcribing sites per cell for both TFF1 and TFF3, as well as in the number of cells showing transcripts compared to the untreated condition (Figure 2- figure supplement 1A-B). Further, TFF1 transcripts spots were more spread out in the nucleus that harboured multiple transcribing sites (Figure 2-figure supplement 1C-D). To check the transcriptional status of TFF1 and TFF3 in the same cell, we plotted the transcript counts from individual cells (Figure 2C). The data suggested that the transcript counts for TFF1 increased at 1h and cells that showed high counts for TFF1 were likely to have low counts for TFF3 and vice-versa. But the RNA counts for TFF1 decreased at 3h and increased for TFF3. Overall, it was evident that the transcriptional profile for these two genes located in the same TAD were negatively correlated as they peaked at different time points. To further confirm this, smFISH using probes targeting both the intronic and exonic transcripts in the same experiment was conducted. Intronic probes represent active transcription, while exonic probes show accumulated mature RNA from past transcriptional events even in the absence of active transcription (Skinner et al., 2016). Representative images from the smFISH experiment probing InTFF1/ExTFF1 and InTFF3/ExTFF3 are depicted (Figure 2D, Figure 2- figure supplement 1E). InTFF1 spots overlapped with ExTFF1 in the nucleus (Figure 2D, Figure 2- figure supplement 1F) and InTFF1 spots were restricted to the nucleus, whereas ExTFF1 were both in the nucleus and the cytoplasm as expected, indicating the specificity of smFISH probes (Figure 2D, Figure 2 - figure supplement 1E). We observed that the exonic transcript counts for both TFF1 and TFF3 increased at 1h compared to uninduced (Figure 2E, F). Strikingly, at 3h, the exonic transcripts for TFF1 continued to increase even while the intronic transcript counts reduced, though the fold increase (1.12 ±0.02) was less compared to that between uninduced and 1h (1.63 ± 0.28), indicating a plateauing of steady state levels. Exonic transcripts for TFF3 at 3h remained comparable to 1h while the intronic counts increased suggesting, TFF1 transcription reduces at 3h, whereas TFF3 expression increases. We reasoned that the ratio of intronic transcript number (InTFF) to exonic transcripts number (ExTFF) should represent the status of transcription as an increase in the number of intronic transcript due to expression would result in a higher ratio while also taking into consideration the number of mature transcripts. As expected, the ratio of intronic transcripts to exonic transcripts also showed that transcription is active at 1h for TFF1 as the ratio is higher compared to uninduced and 3h (Figure 2G, Figure 2- figure supplement 2A). Contrastingly, the ratio was highest at 3h for TFF3 suggesting that active transcription takes place much after the peak of E2-mediated signalling and maximal TFF1 expression (Figure 2H, Figure 2- figure supplement 2B). Indeed, a very recent study has shown that such post-transcriptional splicing can occur for genes and intron dispersal can be expected more commonly for some highly expressed genes (Coté et al., 2023). Therefore, the observation of intronic signal away from the site of transcription as we see for TFF1 and TFF3 is not unexpected (Figure 2A, D, Figure 2-figure supplement 1).

TFF1 enhancer does not change target promoters during signalling time-course

Earlier version: In order to identify the molecular players behind the differential expression of these two genes that are in the same TAD, we asked if the ERα-bound enhancer downstream to the TFF1 gene loops with TFF1 at 1h, and with TFF3 at 3h. This enhancer acutely activates TFF1 at 1h post estrogen stimulation (Saravanan et al., 2020; Ho et al., 2021). We interrogated the looping using enhancer as a viewpoint by 4C-seq in uninduced, 1h and 3h post E2 stimulation. We observed robust interactions between enhancer and TFF1 promoter at 1h post induction, which reduced at 3h. On the other hand, its interaction with the TFF3 promoter exhibited very low counts in uninduced as well as E2-induced conditions at both time points and in both replicates (Fig. 3A, Fig. S4). This suggests that the looping of enhancer potentially induces the expression of TFF1 gene at 1h, and loss of interactions at 3h results in weak TFF1 transcription. However, the lack of interactions between enhancer and TFF3 did not explain the gain of TFF3 expression at 3h post ligand stimulation.

Revised version: In order to identify the molecular players behind the differential expression of these two genes that are in the same TAD, we asked if the ERα-bound enhancer downstream to the TFF1 gene loops with TFF1 at 1h, and with TFF3 at 3h. This enhancer acutely activates TFF1 at 1h post estrogen stimulation (Saravanan et al., 2020; Ho et al., 2021). We interrogated the looping using enhancer as a viewpoint by 4C-seq in uninduced, 1h and 3h post E2 stimulation. We observed robust interactions between enhancer and TFF1 promoter at 1h post induction, which reduced at 3h. On the other hand, its interaction with the TFF3 promoter exhibited very low counts in uninduced as well as E2-induced conditions at both time points and in both replicates (Figure 3A, Figure 3-figure supplement 1). This suggests that the looping of enhancer potentially induces the expression of TFF1 gene at 1h, and loss of interactions at 3h results in weak TFF1 transcription. However, the lack of interactions between enhancer and TFF3 did not explain the gain of TFF3 expression at 3h post ligand stimulation.

Enhancer looping does not account for the differential expression of TFF1 and TFF3 genes

Removed in the revised version: A. 4C-seq plot at TFF1 enhancer viewpoint, the interaction with the promoter is highlighted in yellow. The plot is overlaid with H3K27ac, ERα ChIP signal, and gene annotations.

Removed in the revised version: B. Genome browser snapshot of TFF1 region depicting ERα binding in WT lines. The first, second and third ERα ChIP-seq tracks are from WT cells that are vehicle-treated, E2-1h, and E2-3h, respectively. Blue highlighted regions represent the ΔTFF1e region.

Removed in the revised version: C. The mean RNA numbers for InTFF1 in WT (unshaded) and ΔTFF1e (shaded) MCF7 cells are depicted. The mean of means are shown, and error bars denote SEM from three repeats (n>650, N=3 for each WT and delete line). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: D. The mean RNA numbers for InTFF3 in WT (unshaded) and ΔTFF1e (shaded) MCF7 cells are depicted. The mean of means are shown, and error bars denote SEM from three repeats (n>650, N=3 for each WT and delete line). P-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: E. The mean RNA numbers for InTFF1 and InTFF3 in ΔTFF1e MCF7 cells are depicted. The mean of means are shown, and error bars denote SEM from three repeats (n>880, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Earlier version: F. Ratio of InTFF in WT MCF7 to InTFF in ΔTFF1e MCF7 are depicted. The ratio was obtained by dividing the absolute RNA counts of the WT line by delete lines performed on different days but in the same order (replicate one of WT divided by replicate one of ΔTFF1e). The mean of means are shown, and error bars denote SEM from three repeats (n>650, N=3 for each WT and delete line). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Revised version: A. 4C-seq plot at TFF1 enhancer viewpoint, the interaction with the promoter is highlighted in yellow. The plot is overlaid with H3K27ac, ERα ChIP signal, and gene annotations. B. Genome browser snapshot of TFF1 region depicting ERα binding in WT lines. The first, second and third ERα ChIP-seq tracks are from WT cells that are vehicle-treated, E2-1h, and E2-3h, respectively. Blue highlighted regions represent the ΔTFF1e region. C. The mean RNA numbers for InTFF1 in WT (unshaded) and ΔTFF1e (shaded) MCF7 cells are depicted. The mean of means are shown, and error bars denote SEM from three repeats (n>650, N=3 for each WT and delete line). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. D. The mean RNA numbers for InTFF3 in WT (unshaded) and ΔTFF1e (shaded) MCF7 cells are depicted. The mean of means are shown, and error bars denote SEM from three repeats (n>650, N=3 for each WT and delete line). P-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. E. The mean RNA numbers for InTFF1 and InTFF3 in ΔTFF1e MCF7 cells are depicted. The mean of means are shown, and error bars denote SEM from three repeats (n>880, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. F. Ratio of InTFF in WT MCF7 to InTFF in ΔTFF1e MCF7 are depicted. The ratio was obtained by dividing the absolute RNA counts of the WT line by delete lines performed on different days but in the same order (replicate one of WT divided by replicate one of ΔTFF1e). The mean of means are shown, and error bars denote SEM from three repeats (n>650, N=3 for each WT and delete line). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Earlier version: This could mean that enhancer interaction is critical for TFF1 expression, but is less important for TFF3 expression. Thus, TFF1 expression should be affected more severely upon deletion of the enhancer than TFF3. To test this, we investigated the expression of TFF1 and TFF3 in MCF-7 cells where the enhancer downstream to TFF1 was homozygously deleted using CRISPR-Cas9 (referred to as the ΔTFF1e from hereon) (Fig. 3B) (Saravanan et al., 2021). Using smFISH, we quantified the intronic transcripts for TFF1 and TFF3 in the ΔTFF1e cells, compared to WT cells. We observed that the mean number of TFF1 transcripts was reduced drastically in the ΔTFF1e compared to the WT (Fig. 3C). The reduction was far more substantial (49.23 ±2.6 at 1hr and 40.75 ±7.8 at 3h) for TFF1 than TFF3 (6.03 ±1.8 at 1hr and 7.6 ±0.7 at 3h) (Fig. 3D). The absolute transcript counts for TFF1 and TFF3 in the ΔTFF1e cells have also been shown for clearer visualization (Fig. 3E), as these are obscured when compared to WT. To get a sense of fold changes, we took a ratio of mean transcript counts in WT cells to ΔTFF1e cells at each time point. We observed that the drop in gene transcripts between WT and ΔTFF1e were several folds higher for TFF1 compared to TFF3 (Fig. 3F). This suggests that enhancer deletion has a more robust impact on the transcription of TFF1 compared to TFF3. Nonetheless, TFF3 was also affected even though, it does not loop with the enhancer. This is in accordance with the 4C-seq data where we observed prominent looping between the enhancer and TFF1 but less so with TFF3. These results suggest that the enhancer plays a more important role in the expression of the primary gene while it has less impact on a gene located more distally but within the same TAD during course of signalling.

Revised version: This could mean that enhancer interaction is critical for TFF1 expression, but is less important for TFF3 expression. Thus, TFF1 expression should be affected more severely upon deletion of the enhancer than TFF3. To test this, we investigated the expression of TFF1 and TFF3 in MCF-7 cells where the enhancer downstream to TFF1 was homozygously deleted using CRISPR-Cas9 (referred to as the ΔTFF1e from hereon) (Figure 3B) (Saravanan et al., 2021). Using smFISH, we quantified the intronic transcripts for TFF1 and TFF3 in the ΔTFF1e cells, compared to WT cells. We observed that the mean number of TFF1 transcripts was reduced drastically in the ΔTFF1e compared to the WT (Figure 3C). The reduction was far more substantial (49.23 ±2.6 at 1hr and 40.75 ±7.8 at 3h) for TFF1 than TFF3 (6.03 ±1.8 at 1hr and 7.6 ±0.7 at 3h) (Figure 3D). The absolute transcript counts for TFF1 and TFF3 in the ΔTFF1e cells have also been shown for clearer visualization (Figure 3E), as these are obscured when compared to WT. To get a sense of fold changes, we took a ratio of mean transcript counts in WT cells to ΔTFF1e cells at each time point. We observed that the drop in gene transcripts between WT and ΔTFF1e were several folds higher for TFF1 compared to TFF3 (Figure 3F). This suggests that enhancer deletion has a more robust impact on the transcription of TFF1 compared to TFF3. Nonetheless, TFF3 was also affected even though it does not loop with the enhancer. This is in accordance with the 4C-seq data where we observed prominent looping between the enhancer and TFF1 but less so with TFF3. These results suggest that the enhancer plays a more important role in the expression of the primary gene while it has less impact on a gene located more distally but within the same TAD during the course of signalling.

Levels of ERα in the nucleus dictate the extent of TFF1 and TFF3 inductions

Earlier version: After ruling out enhancer looping as the determinant of differential expression, we looked for other candidates that could regulate the differential gene expression. As discussed above, globally, E2-mediated gene expression is known to peak at 1h after stimulation and drop significantly by 3h (Hah et al., 2011). Upon ligand stimulation, ERα translocate into the nucleus, increasing mean ERα intensity in the nucleus which is high at 1h and then decreases significantly by 3h due to degradation. These changes in intensities have been captured using immunofluorescence for ERα (Saravanan et al., 2020). We tested if the intensities of ERα in individual cells correlate with the expression of TFF1 and TFF3 at 1h and 3h of E2 signalling. Towards this, we combined smFISH with immunofluorescence for ERα. To improve contrast for the ERα signal, we additionally performed a chromatin retention assay to get rid of any chromatin unbound ERα (Fig. 4A). The representative images show that the cells with very high levels of nuclear ERα (blue circle) exhibited low counts of both TFF1 and TFF3. In contrast, the cells with medium levels of ERα (red circle) possessed higher TFF1 than TFF3. Similarly, the cells with the lower levels of ERα (grey circle) showed higher TFF3 expression as compared to TFF1 (Fig. 4A). Histograms depicting the ERα mean intensities across individual cells, showed that the nuclear level of ERα increases post 1h of induction and then goes down at 3h (Fig. 4B, C), similar to TFF1 expression (Fig. 4D).

Revised version: After ruling out enhancer looping as the determinant of differential expression, we looked for other candidates that could regulate the differential gene expression. As discussed above, globally, E2-mediated gene expression is known to peak at 1h after stimulation and drop significantly by 3h (Hah et al., 2011). Upon ligand stimulation, ERα translocates into the nucleus, increasing mean ERα intensity in the nucleus which is high at 1h and then decreases significantly by 3h due to degradation. These changes in intensities have been captured using immunofluorescence for ERα (Saravanan et al., 2020). We tested if the intensities of ERα in individual cells correlate with the expression of TFF1 and TFF3 at 1h and 3h of E2 signalling. Towards this, we combined smFISH with immunofluorescence for ERα. To improve contrast for the ERα signal, we additionally performed a chromatin retention assay to get rid of any chromatin unbound ERα (Figure 4A). The representative images show that the cells with very high levels of nuclear ERα (blue circle) exhibited low counts of both TFF1 and TFF3. In contrast, the cells with medium levels of ERα (red circle) possessed higher TFF1 than TFF3. Similarly, the cells with the lower levels of ERα (grey circle) showed higher TFF3 expression as compared to TFF1 (Figure 4A). Histograms depicting the ERα mean intensities across individual cells, showed that the nuclear level of ERα increases post 1h of induction and then goes down at 3h (Figure 4B, C), similar to TFF1 expression (Figure 4D).

Earlier version: Nuclear levels of the transcription factor ERα dictate the extent of expression of TFF1 and TFF3 genes

Revised version: Nuclear levels of ERα dictate the extent of TFF1 and TFF3 expression

Removed in the revised version: A. Representative images showing smFISH for InTFF1 and InTFF3 in combination with immunofluorescence for ERα (along with chromatin retention assay). The red circle denotes a cell with high ERα and low TFF1 and TFF3, the blue circle denotes a cell with medium ERα, high TFF1, and low TFF3 while the green circle denotes a cell with low ERα and high TFF3. The scale bar is 5 microns.

Removed in the revised version: B. Histogram representing the distribution of ERα mean intensities in cells under uninduced, E2-1hr, and E2-3hr conditions (n=210). Intensities at 1hr are the highest while they shift to the left at 3hr and are lowest in uninduced cells. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another.

Removed in the revised version: C. Cumulative histogram representing the distribution of ERα mean intensities in cells under uninduced, E2-1hr, and E2-3hr conditions. p-values were calculated by the Mann-Whitney test and the significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: D. ERα intensities were sorted into three categories namely low (intensities between 0-450 A.U.) mid (intensities between 450-1200 A.U.) and high (intensities between 1200-2100 A.U.). The mean and SEM of transcript count for InTFF1 in the three categories under uninduced, E2-1hr, and E2-3hr was plotted. Low and mid categories show the highest TFF1 mean. p-values were calculated by the Mann-Whitney test and the significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another.

Earlier version: E. ERα intensities were sorted into three categories namely low (intensities between 0-450 A.U.) mid (intensities between 450-1200 A.U.) and high (intensities between 1200-2100 A.U.). The mean and SEM of transcript count for InTFF3 in the three categories under uninduced, E2-1hr, and E2-3hr was plotted. Low category shows the highest TFF3 mean. p-values were calculated by the Mann-Whitney test and the significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another.

Revised version: A. Representative images showing smFISH for InTFF1 and InTFF3 in combination with immunofluorescence for ERα (along with chromatin retention assay). The red circle denotes a cell with high ERα and low TFF1 and TFF3, the blue circle denotes a cell with medium ERα, high TFF1, and low TFF3 while the green circle denotes a cell with low ERα and high TFF3. The scale bar is 5 micrometres. B. Histogram representing the distribution of ERα mean intensities in cells under uninduced, E2-1hr, and E2-3hr conditions (n=210). Intensities at 1hr are the highest while they shift to the left at 3hr and are lowest in uninduced cells. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another. C. Cumulative histogram representing the distribution of ERα mean intensities in cells under uninduced, E2-1hr, and E2-3hr conditions. p-values were calculated by the Mann-Whitney test and the significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. D. ERα intensities were sorted into three categories namely low (intensities between 0-450 A.U.) mid (intensities between 450-1200 A.U.) and high (intensities between 1200-2100 A.U.). The mean and SEM of transcript count for InTFF1 in the three categories under uninduced, E2-1hr, and E2-3hr was plotted. Low and mid categories show the highest TFF1 mean. p-values were calculated by the Mann-Whitney test and the significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another. E. ERα intensities were sorted into three categories namely low (intensities between 0-450 A.U.) mid (intensities between 450-1200 A.U.) and high (intensities between 1200-2100 A.U.). The mean and SEM of transcript count for InTFF3 in the three categories under uninduced, E2-1hr, and E2-3hr was plotted. Low category shows the highest TFF3 mean. p-values were calculated by the Mann-Whitney test and the significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another. F. 3D plot representing the distribution of ERα, InTFF1, and InTFF3 on a cell-by-cell basis shows that cells with lower levels of ERα show higher counts for InTFF3. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another.

Removed in the revised version: F. 3D plot representing the distribution of ERα, InTFF1, and InTFF3 on a cell-by-cell basis shows that cells with lower levels of ERα show higher counts for InTFF3. This is plotted from one experimental repeat out of three repeats as ERα intensities will vary from one immunofluorescence experiment to another.

Earlier version: To further corroborate this, we parsed the ERα population cells into 3 categories, namely low, medium, and high within the cells imaged at 1 and 3h. We plotted the mean counts of TFF1 in each of these bins at different time points and observed that the mean count was higher in the mid-category (Fig. 4D). While for TFF3, the mean count was significantly higher in the low bin (Fig. 4E). The transcript counts for TFF1 and TFF3 against ERα intensities on a cell-by-cell basis (Fig. 4F), also showed this feature where very high levels of ERα in the nucleus were not conducive to the expression of either gene (Fig. 4F). To test if ERα had a causal role in intensity based expression of TFF1 and TFF3, we increased the levels of ERα by overexpression of ERα-GFP (Fig. S5A). The representative images from the smFISH experiment in cells overexpressing ERα-GFP confirmed that TFF1 and TFF3 were downregulated in transfected cells while these genes were not perturbed in non-transfected cells in the neighbourhood (Fig. S5A,C). Meanwhile, the cells overexpressing ERα-GFP do not show any impairment in the expression of GAPDH (housekeeping gene) (Fig. S5A,B). As another control, we transfected the cells with EGFP-C1 (same backbone as ERα-GFP construct) and observed no effect on TFF1 or TFF3 (Fig. S5D,E). The data suggest that loss of TFF1 and TFF3 expression upon ERα overexpression was not a general effect of transfection stress but rather specific to ERα overexpression. These results indicate that high nuclear levels of ERα can be detrimental to the expression of genes it regulates. This could be due to the widespread condensate formation, which in turn could sequester the transcriptional protein complexes and competitively abrogate transcription across multiple loci. Thus, the global level of ERα in the nucleus can predict the transcriptional status of specific genes. The binding of ERα at 1 and 3h is proportional to its nuclear levels (Fig. 3A), suggesting, its overexpression would lead to more binding in the genome, which is detrimental to gene expression.

Revised version: To further corroborate this, we parsed the ERα population cells into 3 categories, namely low, medium, and high within the cells imaged at 1 and 3h. We plotted the mean counts of TFF1 in each of these bins at different time points and observed that the mean count was higher in the mid-category (Figure 4D). While for TFF3, the mean count was significantly higher in the low bin (Figure 4E). The transcript counts for TFF1 and TFF3 against ERα intensities on a cell-by-cell basis (Figure 4F), also showed this feature where very high levels of ERα in the nucleus were not conducive to the expression of either gene (Figure 4F). To test if ERα had a causal role in intensity based expression of TFF1 and TFF3, we increased the levels of ERα by overexpression of ERα-GFP (Figure 4- figure supplement 1A). The representative images from the smFISH experiment in cells overexpressing ERα-GFP confirmed that TFF1 and TFF3 were downregulated in transfected cells while these genes were not perturbed in non-transfected cells in the neighbourhood (Figure 4- figure supplement 1A, C). Meanwhile, the cells overexpressing ERα-GFP do not show any impairment in the expression of GAPDH (housekeeping gene) (Figure 4-figure supplement 1A, B). As another control, we transfected the cells with EGFP-C1 (same backbone as ERα-GFP construct) and observed no effect on TFF1 or TFF3 (Figure 4- figure supplement 1D, E). The data suggest that loss of TFF1 and TFF3 expression upon ERα overexpression was not a general effect of transfection stress but rather specific to ERα overexpression. These results indicate that high nuclear levels of ERα can be detrimental to the expression of genes it regulates. This could be due to the widespread condensate formation, which in turn could sequester the transcriptional protein complexes and competitively abrogate transcription across multiple loci. Thus, the global level of ERα in the nucleus can predict the transcriptional status of specific genes. The binding of ERα at 1 and 3h is proportional to its nuclear levels (Figure 3A), suggesting its overexpression would lead to more binding in the genome, which is detrimental to gene expression.

Earlier version: LLPS perturbation down-regulates TFF1 but supports TFF3 expression

Revised version: 1,6 HD exposure down-regulates TFF1 but supports TFF3 expression

Earlier version: The data obtained from combined smFISH and ERα immunofluorescence indicates that ERα could be a determining factor in controlling the differential gene expression of TFF1 and TFF3. Existing literature indicates that ERα-mediated condensate plays a role in E2-induced gene expression (Boija et al., 2018; Nair et al, 2019; Sabari et al., 2018; Saravanan et al., 2020). It led us to hypothesize that ERα-mediated condensate at the TFF1 locus could be sequestering all the factors required for active transcription thus preventing activation of the TFF3 locus at the active phase of E2 signalling. To validate this, we treated the cells with 3% 1,6-Hexanediol for 5 minutes, which is known to disrupt LLPS (Gamliel et al., 2022; Kroschwald et al., 2017). Following the treatment, we performed smFISH to look at the transcription of TFF1 and TFF3 in the same cell and observed a dramatic reduction in TFF1 transcripts, whereas a statistically significant increase in TFF3 was noted (Fig. 5A). Since ERα forms condensate only after estrogen stimulation, 1,6-Hexanediol had no effect on TFF1 in the absence of E2 signalling (Fig. S6). Together, these results suggest that the functional loss of TFF1 promoter transcription due to the dissolution of ERα condensate allowed the TFF3 promoter in the neighbourhood to gain access to transcriptional machinery, leading to its upregulation. In order to test the generality of this observation, beyond the TFF1 and TFF3 locus, we divided E2-responsive genes into three categories, low, moderate, and high, based on their expression upon E2-40m stimulation. We observed a significant increase at 40m and down-regulation at 160m when the binding of ERα is reduced in the genome (Fig. 5B). Next, we tested the expression of their nearby upregulated genes at 160m compared to 40m post E2 stimulation. The nearby genes are within a distance of 1Mb from the E2-responsive genes, which is approximately the average size of a TAD. Indeed, we observed a significant upregulation of nearby genes at 160m when the expression of highly induced genes dropped. Additionally, the expression of these neighbouring genes was reduced at the peak of signalling (Fig. 5B), showing that at 40m of signalling, the acute activation of primary genes and sequestration of transcription machinery by these genes leads to the loss of expression of nearby genes.

Revised version: The data obtained from combined smFISH and ERα immunofluorescence indicates that ERα could be a determining factor in controlling the differential gene expression of TFF1 and TFF3. Existing literature indicates that ERα-mediated condensate plays a role in E2-induced gene expression (Boija et al., 2018; Nair et al, 2019; Sabari et al., 2018; Saravanan et al., 2020). It led us to hypothesize that ERα-mediated condensate at the TFF1 locus could be sequestering all the factors required for active transcription thus preventing activation of the TFF3 locus at the active phase of E2 signalling. To test this, we treated the cells with 3% 1,6-Hexanediol for 5 minutes, which is known to disrupt LLPS (Gamliel et al., 2022; Kroschwald et al., 2017). Following the treatment, we performed smFISH to look at the transcription of TFF1 and TFF3 in the same cell and observed a dramatic reduction in TFF1 transcripts, whereas a statistically significant increase in TFF3 was noted (Figure 5A). Since ERα forms condensate only after estrogen stimulation, 1,6-Hexanediol had no effect on TFF1 in the absence of E2 signalling (Figure 5- figure supplement 1). Together, these results suggest that the functional loss of TFF1 promoter transcription potentially due to the dissolution of ERα condensate allowed the TFF3 promoter in the neighbourhood to gain access to transcriptional machinery, leading to its upregulation. In order to test the generality of this observation, beyond the TFF1 and TFF3 locus, we divided E2-responsive genes into three categories, low, moderate, and high, based on their expression upon E2-40m stimulation. We observed a significant increase at 40m and down-regulation at 160m when the binding of ERα is reduced in the genome (Figure 5B). Next, we tested the expression of their nearby upregulated genes at 160m compared to 40m post E2 stimulation. The nearby genes are within a distance of 1Mb from the E2-responsive genes, which is approximately the average size of a TAD. Indeed, we observed a significant upregulation of nearby genes at 160m when the expression of highly induced genes dropped. Additionally, the expression of these neighbouring genes was reduced at the peak of signalling (Figure 5B), showing that at 40m of signalling, the acute activation of primary genes and sequestration of transcription machinery by these genes leads to the loss of expression of nearby genes.

Added in the revised version: 1,6 HD exposure down-regulates TFF1 but supports TFF3 expression

Added in the revised version: A. Mean transcript counts for InTFF1 and InTFF3 in control, and 3% 1,6-Hexanediol treated cells post 30 minutes of E2-induction. The mean of means are shown, and the error bars denote SEM from three repeats (n>880, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. B. Boxplots showing DESeq2 normalized counts for low expressing, moderately expressing, and highly-expressing genes in the vehicle, E2-40m, and E2-160m respectively (left). Boxplots showing DESeq2 normalized counts for genes near low expressing, moderately expressing, and highly expressing genes in the vehicle, E2-40m, and E2-160m, respectively (right). The p-values in the boxplots were calculated using the Wilcoxon rank-sum test. The boxplots depict the minimum, first quartile, median, third quartile, and maximum values, along with outliers. C. Model depicting the signaling under uninduced, E2-1hr, and E2-3hr conditions – First, activation of target gene loci (like TFF1) occurs by ligand-dependent induction. During the active phase (1 hr), liganded ERα binds on enhancer and promoter. Together these elements interact in 3D, manifesting as ERα punctae, which results in robust expression of target genes, and the sequestration of transcriptional machinery. At 3 hrs, as nuclear ERα decreases, the transcriptional machinery becomes available to other nearby promoters leading to increase in gene transcription at those loci.

Removed in the revised version: ERα mediated phase separation could govern the differential expression of TFF1 and TFF3 genes

Removed in the revised version: A. Mean transcript counts for InTFF1 and InTFF3 in control, and 3% 1,6-Hexanediol treated cells post 30 minutes of E2-induction. The mean of means are shown, and the error bars denote SEM from three repeats (n>880, N=3). p-values were calculated by Student’s two-tailed unpaired t-test, and the significance is represented as: *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05.

Removed in the revised version: B. Boxplots showing DESeq2 normalized counts for low expressing, moderately expressing, and highly-expressing genes in the vehicle, E2-40m, and E2-160m respectively (left). Boxplots showing DESeq2 normalized counts for genes near low expressing, moderately expressing, and highly expressing genes in the vehicle, E2-40m, and E2-160m, respectively (right). The p-values in the boxplots were calculated using the Wilcoxon rank-sum test. The boxplots depict the minimum, first quartile, median, third quartile, and maximum values, along with outliers.

Removed in the revised version: C. Model depicting the signaling under uninduced, E2-1hr, and E2-3h conditions-1. Activation of TFF1 locus by ligand induction. 2. During the active phase, liganded ERα binds on enhancer and promoter. Together these in 3D interact, manifesting as ERα puncta, which results in robust expression of target genes, 3. The phase-separated condensate leads to the sequestration of transcriptional machinery. 4. This favours intra-TAD enhancer: promoter interactions to facilitate robust gene transcription. Upon ERα degradation at 3h, these condensates disappear, and transcriptional machinery becomes available to other promoters leading to their somewhat increased gene transcription.

Discussion

Earlier version: smFISH allowed us to simultaneously capture allele level transcription of two genes, TFF1 and TFF3, located in the same TAD, at single cell level during the peak and fall of signalling. We were able to capture the anti-correlated expression of these two genes, revealing intricate regulatory feedback between acutely activated enhancer-dependent gene, TFF1, which caused the dysregulation of the non-enhancer targeted gene, TFF3, within the same TAD (Fig. 5B).

Revised version: smFISH allowed us to simultaneously capture allele level transcription of two genes, TFF1 and TFF3, located in the same TAD, at single cell level during the peak and fall of signalling. We were able to capture the anti-correlated expression of these two genes, revealing intricate regulatory feedback between acutely activated enhancer-dependent gene, TFF1, which caused the dysregulation of the non-enhancer targeted gene, TFF3, within the same TAD (Figure 5B).

Earlier version: Our data suggests that while condensate formation on enhancer allows robust activation of the target gene, however, it negatively impacts the expression of other neighbouring genes potentially due to the sequestration of transcriptional machinery from these genes. When the condensates dissolve, the locally enriched transcriptional machinery is available to the other loci, potentially allowing for increased transcription of these genes, that are not the direct target of enhancers (Fig. 5C). Forced dissolution of condensates allows the expression of non-enhancer target genes in the cells, although at low levels. The data also suggests that, globally, the genes next to ligand-induced genes are upregulated only at the fall of signalling as an indirect consequence of excess polymerase availability in the neighbourhood.

Revised version: Our data suggests that while potential condensate formation on enhancer allows robust activation of the target gene, however, it negatively impacts the expression of other neighbouring genes possibly due to the sequestration of transcriptional machinery from these genes. When the condensates dissolve, the locally enriched transcriptional machinery is available to the other loci, allowing for increased transcription of these genes that are not the direct target of enhancers (Figure 5C). Forced dissolution of condensates by 1,6 HD allows the expression of non-enhancer target genes in the cells, although at low levels. The data also suggests that, globally, the genes next to ligand-induced genes are upregulated only at the fall of signalling as an indirect consequence of excess polymerase availability in the neighbourhood.

Non-enhancer target genes are also regulated indirectly in the same TAD

Earlier version: ERα binding strength increases at the TFF1 enhancer at 1h of ligand stimulation (Saravanan et al., 2020). Such estrogen-induced binding of ERα leads to robust ligand-induced activation of TFF1 gene. While TFF3 promoter remains unbound by ERα (Fig. 3A) and does not interact with the TFF1 enhancer throughout the course of signalling. These data suggest that TFF3 expression is both ERα and enhancer independent. Despite non-dependence, TFF3 expression was mutually exclusive to TFF1 expression and was negatively impacted in the absence of the enhancer. The latter could be due to the enhancer-mediated repositioning of the entire TAD that benefitted the expression of TFF3 and not due to the direct interaction between the enhancer and the TFF3 gene.

Revised version: ERα binding strength increases at the TFF1 enhancer at 1h of ligand stimulation (Saravanan et al., 2020). Such estrogen-induced binding of ERα leads to robust ligand-induced activation of TFF1 gene. While TFF3 promoter remains unbound by ERα (Figure 3A) and does not interact with the TFF1 enhancer throughout the course of signalling. These data suggest that TFF3 expression is both ERα and enhancer independent. Despite non-dependence, TFF3 expression was mutually exclusive to TFF1 expression and was negatively impacted in the absence of the enhancer. The latter could be due to the enhancer-mediated repositioning of the entire TAD that benefitted the expression of TFF3 and not due to the direct interaction between the enhancer and the TFF3 gene.

Earlier version: The mean expression of TFF1 was manyfold higher than TFF3, thus explaining its dependence on an active upstream enhancer. The enhancer of TFF1 interacts with the gene at the peak of signalling and shows some loss of interaction at 3h when signalling recedes. As stated above, the enhancer does not interact with TFF3, however, the enhancer forms loop with TFF3 upon deletion of the TFF1 promoter (Oh et al., 2021). The ERα condensates are formed at the TFF1 locus upon signalling (Saravanan et al., 2020), which is reduced upon enhancer deletion, suggesting its causal role in condensate formation. Corroborating with this, the enhancer of TFF1 engages in the recruitment of megadalton complex (Liu et al., 2014) in MCF-7 cells, which could facilitate phase-separation (Nair et al., 2019, Saravanan et al., 2020). The cells that expressed TFF1 did not express TFF3 as efficiently, suggesting Pol2 was potentially sequestered from the TFF3 promoter. However, at 3h when liganded ERα degrades leading to dissolution of condensates on enhancer, TFF3 expression increased, suggesting a local redistribution of active transcription machinery to other genes within the same chromatin domain. Even at 1h, when condensates were perturbed by 1,6 HD, the expression of TFF3 increased indicating, the access to Pol2 pool by TFF3 promoter. The data explains the negative correlation between TFF1 and TFF3 expression at the peak of signalling and robust activation of TFF3 at 3h when TFF1 expression decreased at single cell level. The presence of significant numbers of TFF1 nascent RNA in the nucleus in our data could also be due to a time lapse between transcription and splicing, as seen for the genes that are early responders during signaling (Zambrano et al., 2020).

Revised version: The mean expression of TFF1 was manyfold higher than TFF3, thus explaining its dependence on an active upstream enhancer. The enhancer of TFF1 interacts with the gene at the peak of signalling and shows some loss of interaction at 3h when signalling recedes. As stated above, the enhancer does not interact with TFF3, however, the enhancer forms a loop with TFF3 upon deletion of the TFF1 promoter (Oh et al., 2021). The ERα condensates are formed at the TFF1 locus upon signalling (Saravanan et al., 2020), which is reduced upon enhancer deletion, suggesting its causal role in condensate formation. Corroborating with this, the enhancer of TFF1 engages in the recruitment of megadalton complex (Liu et al., 2014) in MCF-7 cells, which could facilitate phase-separation (Nair et al., 2019, Saravanan et al., 2020). The cells that expressed TFF1 did not express TFF3 as efficiently, suggesting Pol2 was potentially sequestered from the TFF3 promoter. However, at 3h when liganded ERα degrades leading to dissolution of potential condensates on enhancer, TFF3 expression increased, suggesting a local redistribution of active transcription machinery to other genes within the same chromatin domain. Even at 1h, when condensates were perturbed by 1,6 HD, the expression of TFF3 increased indicating the access to Pol2 pool by TFF3 promoter. The data explains the negative correlation between TFF1 and TFF3 expression at the peak of signalling and robust activation of TFF3 at 3h when TFF1 expression decreased at single cell level. The presence of significant numbers of TFF1 nascent RNA in the nucleus in our data could also be due to a time lapse between transcription and splicing, as seen for the genes that are early responders during signaling (Zambrano et al., 2020).

High and low levels of ERα were detrimental to enhancer mediated activation of TFF1

Earlier version: We observe that TFF1 expression is the highest in cells with medium levels of ERα, while it shows low expression in cells with low or high levels of ERα (Fig.4D, F). This points to an inverted U-shaped response of TFF1 to ERα, which is a nonmonotonic response to estrogen, as has been shown (Lebedeva et al., 2012; Vandenberg et al., 2012). A nonmonotonic response can be described as an inverted U-shaped curve with maximal response occurring in the middle and reducing substantially at very low or very high doses (Vandenberg et al., 2012). Further, we observe that for TFF3, the maximal expression occurs at low levels of ERα (Fig.4E, F). This could be indicative of a response such that a low concentration is enough to meet the threshold for expression beyond which all other higher doses attenuate the expression (Lebedeva et al., 2012; Vandenberg et al., 2012). This could also suggest that ERα levels have a varied effect on different genes and their expression depending on the extent of reliance of the gene on ERα for its expression. This phenomenon, called the prozone effect, is characterized by the inhibition of multimeric complexes when one of the constituents of the complex are present at very high concentrations as has been shown for TFF1 (Lebedeva et al., 2012). In the context of condensates on TFF1 gene, we speculate that very low levels of TF may not reach or exceed the critical concentration required for LLPS (for TFF1), but this poor transcriptional state of TFF1 could be favourable for TFF3. Alternatively, a very high concentration of the TF could result in the formation of highly dense homogenous condensates, effectively reducing transcription in a situation akin to the prozone effect (Chong et al., 2022; Lebedeva et al., 2012; Ryu et al., 2024).

Revised version: We observe that TFF1 expression is the highest in cells with medium levels of ERα, while it shows low expression in cells with low or high levels of ERα (Figure 4D, F). This points to an inverted U-shaped response of TFF1 to ERα, which is a nonmonotonic response to estrogen, as has been shown (Lebedeva et al., 2012; Vandenberg et al., 2012). A nonmonotonic response can be described as an inverted U-shaped curve with maximal response occurring in the middle and reducing substantially at very low or very high doses (Vandenberg et al., 2012). Further, we observe that for TFF3, the maximal expression occurs at low levels of ERα (Figure 4E, F). This could be indicative of a response such that a low concentration is enough to meet the threshold for expression beyond which all other higher doses attenuate the expression (Lebedeva et al., 2012; Vandenberg et al., 2012). This could also suggest that ERα levels have a varied effect on different genes and their expression depending on the extent of reliance of the gene on ERα for its expression. This phenomenon, called the prozone effect, is characterized by the inhibition of multimeric complexes when one of the constituents of the complex are present at very high concentrations as has been shown for TFF1 (Lebedeva et al., 2012). In the context of condensates on TFF1 gene, we speculate that very low levels of TF may not reach or exceed the critical concentration required for LLPS (for TFF1), but this poor transcriptional state of TFF1 could be favourable for TFF3. Alternatively, a very high concentration of the TF could result in the formation of highly dense homogenous condensates, effectively reducing transcription in a situation akin to the prozone effect (Chong et al., 2022; Lebedeva et al., 2012; Ryu et al., 2024).

Earlier version: Interestingly, the levels of ERα do not determine the expression of other E2 target genes like GREB1 and MYC (Stossi et al., 2020); however, we observed that overexpression of ERα causes downregulation of TFF1 (Fig. S5A) and this is in agreement with what has been observed for the expression status of TFF1 in cell line with high ERα compared to cell lines with low ERα (Lebedeva et al., 2012).

Revised version: Interestingly, the levels of ERα do not determine the expression of other E2 target genes like GREB1 and MYC (Stossi et al., 2020); however, we observed that overexpression of ERα causes downregulation of TFF1 (Figure 4- figure supplement 1A) and this is in agreement with what has been observed for the expression status of TFF1 in cell line with high ERα compared to cell lines with low ERα (Lebedeva et al., 2012).

ER is known to be expressed in nearly 70% of breast cancers, and many of these are treated using estrogen receptor modulators and degraders (Arnesen et al., 2021). Therefore, understanding the effect of ERα on the expression of estrogen-responsive genes depending on whether they are target genes or non-target genes, is of therapeutic importance as several non-target genes have been shown to be differentially regulated in ER-mutant breast cancer cells (Arnesen et al., 2021).

Earlier version: Together, our study underscores the indirect effects of ligand-induced chronic gene transcription that is dependent on enhancer activation and phase separation within a TAD.

Revised version: Together, our study underscores the indirect effects of ligand-induced chronic gene transcription that is dependent on enhancer activation and possible TF-condensate within a TAD.

Materials and Methods

Cell culture

Earlier version: MCF-7 cells (WT scramble control and ΔTFF1e were generated in Saravanan et al., 2020) cells were cultured in non-stripping media consisting of DMEM (Gibco, 12100-046) supplemented with 10% FBS (Gibco, 16000-044) and 1% Penicillin-Streptomycin-Glutamine (Gibco, 10378-016) in a 5% CO2 humidified incubator at 37 °C (unstripped MCF-7). Cells were passaged and seeded into glass bottom dishes in non-stripping media for 24h and allowed to reach 60–80% confluency. These cells were then hormone-stripped for three days in stripping media containing phenol red-free DMEM (Gibco, 21063-029) supplemented with 5% charcoal–stripped FBS(Gibco,12676029) and maintained in the humidified incubator at 37 °C (stripped MCF-7).

Revised version: MCF-7 cells (original line procured from ATCC; WT scramble control and ΔTFF1e were generated in Saravanan et al., 2020) were cultured in non-stripping media consisting of DMEM (Gibco, 12100-046) supplemented with 10% FBS (Gibco, 16000-044) and 1% Penicillin-Streptomycin-Glutamine (Gibco, 10378-016) in a 5% CO2 humidified incubator at 37 °C (unstripped MCF-7). Cells were passaged and seeded into glass bottom dishes in non-stripping media for 24h and allowed to reach 60–80% confluency. These cells were then hormone-stripped for three days in stripping media containing phenol red-free DMEM (Gibco, 21063-029) supplemented with 5% charcoal–stripped FBS (Gibco,12676029) and maintained in the humidified incubator at 37 °C (stripped MCF-7).

In order to induce the estrogen transcriptional responses, on the third day cells were treated with β-estradiol (E2758, Sigma-Aldrich) at 100nM concentration for various periods as mentioned in the respective Figures. For untreated control, cells were either treated with equal microliters of ethanol on the third day or with ERα inhibitor ICI182780 (1047, Tocris Biosciences) at 100nM concentration for 24 h after two days of stripping.

Earlier version: For perturbation of LLPS in E2-induced condition, cells were incubated with media containing 3% 1,6-Hexanediol for 5 minutes (after 30 minutes of E2 treatment) followed by removal of 1,6-Hexanediol and recovery for 25 mins in E2 containing media.

Revised version: For perturbation of transcriptional condensates in E2-induced condition, cells were incubated with media containing 3% 1,6-Hexanediol for 5 minutes (after 30 minutes of E2 treatment) followed by removal of 1,6-Hexanediol and recovery for 25 mins in E2 containing media.

For experiments with overexpression of EGFP- ERα or EGFP, cells were transfected 24 hours before E2 induction using X-tremeGENE™ HP DNA Transfection Reagent (XTGHP-RO) in stripping media.

Probe for smFISH

Earlier version: Probes were designed, custom-made, and tagged with indicated fluorophores (Quasar 570 or Cal 610) by Biosearch Technologies (https://www.biosearchtech.com/products/rna-fish). Probes were made to target the intronic or exonic regions of the TFF1 and TFF3 genes. Probe sequences can be found in supplementary table 1-4 in supplementary materials.

Revised version: Probes were designed, custom-made, and tagged with indicated fluorophores (Quasar 570 or Cal Fluor 610) by Biosearch Technologies (https://www.biosearchtech.com/products/rna-fish). Probes were made to target the intronic or exonic regions of the TFF1 and TFF3 genes. Probe sequences can be found in supplementary file 1a-d in supplementary materials.

Chromatin retention assay and smFISH

For smFISH experiments without chromatin retention assay cells, the cells were treated with estradiol or vehicular control for indicated times. Following this, the cells were washed with nuclease-free 1X PBS (Ambion, AM9624) twice. This was followed by fixation using 4% paraformaldehyde (PFA, Sigma, P6148) in 1X PBS for 10 minutes at room temperature. The fixative was removed and cells were washed twice with 1X PBS. The cells were permeabilized using 70% ethanol at 4°C overnight. On the next day, the permeabilizing agent was removed and the cells were washed twice with 1X PBS. The cells were washed once with a wash buffer (10% Formamide (Ambion, 9342) and 2X SSC (Ambion, AM9763) in nuclease-free water) for 5 minutes at room temperature. The wash buffer was aspirated and the cells were incubated with a hybridization mix (100 μl of mix contains 10% formamide, 89 μl hybridization buffer, and 1 μl each of indicated probes) overnight at 37°C in a humid chamber. The next day, the solution was removed and cells were washed with the wash buffer at 37°C for 30 minutes. To stain the nuclei, DAPI (Invitrogen, D1306; 2μg/ml) in wash buffer was added to the cells and incubated for 10 minutes at 37°C. The cells were then washed with 2X SSC for 5 minutes at 4°C. The solution was removed and the cells were covered with a few drops of the mountant Vectashield (Vector Labs). The plates were imaged after at least 1h.

For smFISH experiments with chromatin retention assay cells were treated with estradiol or vehicular control for indicated times. Following this, the cells were treated with CSK buffer (consisting of 10 mM PIPES buffer, 100 mM NaCl, 3 mM MgCl2, 300 mM sucrose, and 0.7% Triton-X 100) for 15 minutes at room temperature. After this, the cells were washed with nuclease-free 1X PBS twice. This was followed by fixation using 4% paraformaldehyde in 1X PBS for 10 minutes at room temperature. The fixative was removed and cells were washed twice with 1X PBS. The cells were permeabilized using 0.3% Triton-X 100 (Sigma, T8787) in 1X PBS for 10 minutes at room temperature. The permeabilizing agent was aspirated and the cells were washed twice with 1X PBS. This was followed by a washing step using the wash buffer for 5 minutes at room temperature. Following this, the cells were incubated with the hybridization mix to which the antibody of interest was also added at appropriate dilution. The next day, the hybridization mix containing probes and antibodies was removed. Cells were washed twice with 1X PBS. Then the cells were incubated with 1X PBS containing a secondary antibody against the antibody of interest at room temperature for 2h. This was followed by incubation with the wash buffer at 37°C for 30 minutes. To stain the nuclei, DAPI in wash buffer was added to the cells and incubated for 10 minutes at 37°C. The cells were then washed with 2X SSC for 5 minutes at 4°C. The solution was removed and the cells were covered with a few drops of the mountant Vectashield. The plates were imaged after at least 1h.

Antibody staining/ immunofluorescence

The primary antibody against ERα (Santa Cruz, sc8002(F10)) was added to the hybridization mix at the dilution of 1:400 and incubated overnight at 37°C. The next day, an anti-mouse Alexa Fluor™ 488 secondary antibody (Invitrogen, A11029) was added to 1X PBS at the dilution of 1:1000. Cells were incubated for 2h followed by the continuation of the smFISH protocol as indicated above.

Image acquisition

Earlier version: The plates were imaged on an Olympus IX83 inverted widefield fluorescence microscope with a Retiga 6000 CCD monochrome camera (QImaging). The images were acquired using a 60X, 1.42 N.A. oil immersion objective or a 100X, 1.4 N.A. oil immersion objective. The z-step size was 0.3μm and 35 slices were acquired. The resolution at which the images were acquired is 2752 x 2208. Narrow band-pass filters were used to distinguish the signal from Quasar 570 and Cal 610 labeled probes (ChromaTechnology-49309 and 49310).

Revised version: The plates were imaged on an Olympus IX83 inverted widefield fluorescence microscope with a Retiga 6000 CCD monochrome camera (QImaging). The images were acquired using a 60X, 1.42 N.A. oil immersion objective or a 100X, 1.4 N.A. oil immersion objective. The z-step size was 0.3μm and 35 slices were acquired. The resolution at which the images were acquired is 2752 x 2208. Narrow band-pass filters were used to distinguish the signal from Quasar 570 and Cal Fluor 610 labeled probes (ChromaTechnology- 49309 and 49310).

Image analysis and representation

The images of mRNA channels were subjected to rolling ball background subtraction across the entire Z-stack using Fiji. For representative images, the stacks were Z-projected in Fiji. Transcripts were counted using either Rajlabimagetools (courtesy of Arjun Raj lab) or an in-house MATLAB (Mathworks) script based on earlier work (Femino et al., 1998; Raj et al., 2008). Briefly, the code uses a nuclear mask to count the intron-containing RNA present in the nucleus (InTFF), and another mask for the whole cell to count the mature mRNA, detected by exonic probes (ExTFF), that are present in the nucleus as well as the cytoplasm. The mean intensity for the ERα immunofluorescence channel was also quantified using the MATLAB script on a cell-by-cell basis using the nuclear mask.

Graphing and Statistics

The graphs were plotted using Python 3, MATLAB, and Origin Pro and edited using Adobe Photoshop. To perform a student’s t-test for data combined from three repeats, GraphPad (https://www.graphpad.com/quickcalcs/ttest1/) was used. Non-parametric tests were performed in case of single cell data. For this Mann-Whitney test was used. Significance is represented as : *** denotes p < 0.001,** denotes p < 0.01,* denotes p < 0.05, ns denotes p > 0.05. The specific test used has been mentioned in the relevant figure legends.

Circular Chromatin Conformation Capture-seq

Earlier version: 4C was performed as per the protocol described in van de Werken et al., 2012 with minor variations. MCF7 cells were fixed with fresh formaldehyde (1.5%) and quenched with glycine (125mM) followed by washes with ice-cold 1XPBS (2X) and scraped, pelleted, and stored at - 800C. Lysis buffer [Tris-Cl pH 8.0 (10mM), NaCl (10mM), NP-40 (0.2%), PIC (1X)] was added to the pellets and homogenized by Dounce homogenizer (20 stroked with pestle A followed by pestle B). The 3C digestion was performed with DpnII (200 units, NEB) and ligation was performed by T4 DNA ligase and ligation mix [Triton X-100 (1%), 1× Ligation buffer (10X Ligation buffer-Tris-Cl pH 7.5 (500mM), MgCl2 (100mM), DTT (100mM), BSA (0.105mg/ml), ATP (1.05mM)]. The ligated samples were purified by PCI and subjected to ethanol precipitation. The pellet was eluted in 1X TE (pH 8.0) to obtain the 3C library. The 4C digestion was performed by NlaIII (50 units, NEB), and the samples were ligated, purified, and precipitated similar to the 3C library to obtain the 4C library. The 4C library was subjected to RNaseA treatment and purified with the QIAquick PCR purification kit. The concentration of the library was then measured by Nanodrop and subjected to PCRs using the oligos for the enhancer viewpoint. The samples were next PCR purified using the same kit and subjected to next-generation sequencing with Illumina HiSeq2500/ NOVA seq. The 4C oligos are listed in supplementary table 5.

Revised version: 4C was performed as per the protocol described in van de Werken et al., 2012 with minor variations. MCF7 cells were fixed with fresh formaldehyde (1.5%) and quenched with glycine (125mM) followed by washes with ice-cold 1XPBS (2X) and scraped, pelleted, and stored at - 800C. Lysis buffer [Tris-Cl pH 8.0 (10mM), NaCl (10mM), NP-40 (0.2%), PIC (1X)] was added to the pellets and homogenized by Dounce homogenizer (20 stroked with pestle A followed by pestle B). The 3C digestion was performed with DpnII (200 units, NEB) and ligation was performed by T4 DNA ligase and ligation mix [Triton X-100 (1%), 1x Ligation buffer (10X Ligation buffer- Tris-Cl pH 7.5 (500mM), MgCl2 (100mM), DTT (100mM), BSA (0.105mg/ml), ATP (1.05mM)]. The ligated samples were purified by PCI and subjected to ethanol precipitation. The pellet was eluted in 1X TE (pH 8.0) to obtain the 3C library. The 4C digestion was performed by NlaIII (50 units, NEB), and the samples were ligated, purified, and precipitated similar to the 3C library to obtain the 4C library. The 4C library was subjected to RNaseA treatment and purified with the QIAquick PCR purification kit. The concentration of the library was then measured by Nanodrop and subjected to PCRs using the oligos for the enhancer viewpoint. The samples were next PCR purified using the same kit and subjected to next-generation sequencing with Illumina HiSeq2500/ NOVA seq. The 4C oligos are listed in supplementary file 1e.

4C Data Analysis

Earlier version: The sequenced reads in fastq file were demultiplexed by matching the appropriate primer sequences for each condition without allowing for any mismatches. Demultiplexed reads were processed using 4cseqpipe software. Restriction site tracks were created for the hg38 human genome by mentioning the restriction sites of the first cutter as GATC and that of the second cutter as CATG. The phred scores of demultiplexed reads were changed to phred64 format and the fastq files were converted into raw format. Further, valid 4C reads were mapped to the generated restriction site tracks. Unique fragment ends/non-unique fragment ends were used. The mapped reads were normalized and near cis domainograms at a maximum height of 0.1 were created by using the truncated log mean statistic with a trend resolution of 1kb for the genomic region chr11:42300000-42400000.

Revised version: The sequenced reads in fastq file were demultiplexed by matching the appropriate primer sequences for each condition without allowing for any mismatches. Demultiplexed reads were processed using 4cseqpipe software. Restriction site tracks were created for the hg38 human genome by mentioning the restriction sites of the first cutter as GATC and that of the second cutter as CATG. The phred scores of demultiplexed reads were changed to phred64 format and the fastq files were converted into raw format. Further, valid 4C reads were mapped to the generated restriction site tracks. Unique fragment ends/non-unique fragment ends were used. The mapped reads were normalized and near cis domainograms at a maximum height of 0.1 were created by using the truncated log mean statistic with a trend resolution of 1kb for the genomic region chr11:42300000-42400000. Data given in (Fig 3A_source data 1 file).

RNA Isolation, cDNA synthesis and PCR

Earlier version: Cells were lysed in 1 ml of Trizol (Thermo Fisher Inc.). 200 ul chloroform was added to the sample, briefly vortexed and centrifuged at 12K rpm for 12 mins. The aqueous phase was carefully collected and transferred to the fresh tube. 1 volume of isopropanol was added to the sample and incubated at room temperature for 10 mins to precipitate the RNA. The samples were centrifuged at 12K rpm for 12 mins, supernatant was discarded without disturbing the pellet. The pellet obtained was washed with 75% ethanol. Pellet was air dried and dissolved in RNase free water. RNA obtained was treated with ezDNase (Invitrogen) to remove the traces of contaminating DNA. 1μg of RNA was used for each cDNA synthesis reaction by Superscript IV (Invitrogen) and random hexamers as per manufacturer’s recommendation. The CFX96 touch (Biorad) real time PCR was used for qRT-PCRs. The fold changes were calculated by the ΔΔCt method and individual expression data was normalized to GAPDH mRNA. The p-values were calculated by Student’s unpaired two-tailed t-test from independent four biological replicates. qRT-PCR primers are listed in Supplementary Table 6.

Revised version: Cells were lysed in 1 ml of Trizol (Thermo Fisher Inc.). 200 ul chloroform was added to the sample, briefly vortexed and centrifuged at 12K rpm for 12 mins. The aqueous phase was carefully collected and transferred to the fresh tube. 1 volume of isopropanol was added to the sample and incubated at room temperature for 10 mins to precipitate the RNA. The samples were centrifuged at 12K rpm for 12 mins, supernatant was discarded without disturbing the pellet. The pellet obtained was washed with 75% ethanol. The pellet was air dried and dissolved in RNase free water. RNA obtained was treated with ezDNase (Invitrogen) to remove the traces of contaminating DNA. 1μg of RNA was used for each cDNA synthesis reaction by Superscript IV (Invitrogen) and random hexamers as per manufacturer’s recommendation. The CFX96 touch (Biorad) real time PCR was used for qRT-PCRs. The fold changes were calculated by the ΔΔCt method and individual expression data was normalized to GAPDH mRNA. The p-values were calculated by Student’s unpaired two-tailed t-test from independent four biological replicates. qRT-PCR primers are listed in Supplementary file 1f.

GRO-Seq analysis

Fastq files from GEO accession number GSE43836 were downloaded from European Nucleotide Archive. Reads with base quality <20 in a sliding window of 4 bases and with a length of <36 were removed using Trimmomatic 0.39. Trimmed reads were aligned using bowtie2 2.5.1 with default parameters. Duplicate reads from the alignment files were removed using samtools 1.16.1. De-duplicated aligned reads were assigned in a strand-specific manner to transcript feature of the hg38.ncbiRefSeq.gtf.gz file by allowing multi-mapping reads and considering the largest overlap in case of overlapping features using featureCounts v2.0.3. CPM normalised strand-specific signal files with a bin size of 1 bp were generated using bamCoverage 3.5.1. Differential gene expression analysis was performed with default parameters using Deseq2 1.36.0. Upregulated genes were defined as those with adjusted p-value <0.05 and log2(FC) >1 in their respective conditions. Top, Middle and bottom 10% of upregulated genes based on their base mean value in E2-40m vs VEH condition were subsetted as highly expressing, moderately expressing and low expressing upregulated genes. The closest upregulated genes in E2-160m vs E2-40m near highly expressing, moderately expressing and low expressing upregulated genes in E2-40m vs VEH condition were identified using bedtools closest v2.30.0. Only genes within a genomic distance of 1Mb were considered. Log2 transformed DEseq2 normalised counts with the addition of the arbitrary value 1 was used to compare the gene expression trends across time points and categories of upregulated genes. All plots were generated using R 4.2.2.